// intro2.rs
//
// Make the code print a greeting to the world.
//
// Execute `rustlings hint intro2` or use the `hint` watch subcommand for a
// hint.

// I AM NOT DONE

fn main() {
    println!("Hello there!"); // On remplace printline! par println!. On ajoute aussi ; a la fin de la ligne pour marqué la fin de notre commande
}
