// variables5.rs
//
// Execute `rustlings hint variables5` or use the `hint` watch subcommand for a
// hint.

// I AM NOT DONE

fn main() {
    let number = "T-H-R-E-E"; // don't change this line
    println!("Spell a Number : {}", number);
    let number = 3; // don't rename this variable || On ajoute "let" au début afin d'initialiser la variable number avec pour type i32 (integer)
    println!("Number plus two is : {}", number + 2);
}
